# Your Grades and GPA at Illinois!

In this lab, you will explore Wade's famous GPA dataset and find out about the easiest, hardest, and most popular courses at Illinois!

The dataset you are using is this lab is actually the same dataset that Prof Wade used to make the GPA visualizations :) Check this out here if you haven't seen it already: [https://waf.cs.illinois.edu/discovery/grade_disparity_between_sections_at_uiuc/](https://waf.cs.illinois.edu/discovery/grade_disparity_between_sections_at_uiuc/)

# Part 0 | Your Group!

Edit the next Python cell to add information about who you're working within your lab section:

```python
# First, meet your CAs and TA if you haven't already!
# ...first name is enough, we'll know who they are! :)
ta_name = ""
ca1_name = ""
ca2_name = ""


# Also, make sure to meet your team for this lab! Find out their name, what major they're in,
# and learn something new about them that you never knew before!
partner1_name = ""
partner1_netid = ""
partner1_major = ""

partner2_name = ""
partner2_netid = ""
partner2_major = ""

partner3_name = ""
partner3_netid = ""
partner3_major = ""
```

# Part 1 | Basic data manipulation

## Puzzle 1.1: Loading the dataset

For the first part of our lab, we'll be using the GPA dataset, which should be automatically downloaded when you fetch this notebook. 

In the following cell, load the pandas library and read the dataset into `df`.

<!-- grade_id: cell-51fef9428e349da7, points: 0, grade: false, solution: true -->
```python

```

## Puzzle 1.2: Adding helper columns

Add two new columns, `"Total Students"` and `"Average GPA"`, that compute the total number of students in each course and the average GPA for each course, respectively.

You should not include students with a `"W"` in either column. The grade point scale can be found at https://registrar.illinois.edu/courses-grades/explanation-of-grades/.

Hint (double-click this cell to see it): <span style="color:#ffffff00">use the dot() function; see https://stackoverflow.com/questions/18419962/how-to-compute-weighted-sum-of-all-elements-in-a-row-in-pandas for some examples.</span>

<!-- grade_id: cell-1043285927ac8a85, points: 0, grade: false, solution: true -->
```python
weight = [4, 4, 3.67, 3.33, 3, 2.67, 2.33, 2, 1.67, 1.33, 1, 0.67, 0]
grade = ["A+", "A", "A-", "B+", "B", "B-", "C+", "C", "C-", "D+", "D", "D-", "F"]
df["Total Students"] = ...
df["Average GPA"] = ...
```

<!-- grade_id: cell-46729cc74c213e52, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.
assert(df[
  (df["YearTerm"] == "2010-su") & \
  (df["Subject"] == "STAT") & \
  (df["Number"] == 410) \
]["Total Students"].sum() == 31), "It appears that your \"Total Students\" column is incorrect."
assert(df[
  (df["YearTerm"] == "2018-fa") & \
  (df["Subject"] == "CS") & \
  (df["Number"] == 225) \
]["Total Students"].sum() == 697), "It appears that your \"Total Students\" column is incorrect."
assert( abs(df[
  (df["YearTerm"] == "2018-fa") & \
  (df["Subject"] == "STAT") & \
  (df["Number"] == 100) \
]["Average GPA"].mean() - 3.247) < 0.01), "It appears that your \"Average GPA\" column is incorrect."
## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

# Part 2 | Finding the easiest and hardest courses

## Puzzle 2.1: Hardest Courses in the Dataset

Based on the `"Average GPA"` column, find the **30 hardest courses** at Illinois, i.e., the 30 courses with the lowest "Average GPA". Store the results in `df_hard_course`.

Hint (double-click this cell to see it): <span style="color:#ffffff00">use the nsmallest() function</span>

<!-- grade_id: cell-13f6d92b7a658526, points: 0, grade: false, solution: true -->
```python
df_hard_course = ...
df_hard_course
```

## Puzzle 2.2: "Level" of hardest courses

Find the average `"Number"` of the **30 hardest courses** at Illinois. Store the result as `df_hard_course_avg_no`.

<!-- grade_id: cell-a7106dca7694711c, points: 0, grade: false, solution: true -->
```python
df_hard_course_avg_no = ...
df_hard_course_avg_no
```

<!-- grade_id: cell-46729cc74c213e52, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2.1 and 2.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell(s), make changes, and RE-RUN your code and then this cell.
assert(abs(df_hard_course_avg_no - 4853/30) < 0.0001), "It appears that your average \"Number\" is incorrect. Please check both Puzzle 2.1 and Puzzle 2.2."
## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

## Puzzle 2.3: Are the hardest courses really "hard"?

❓ **Individual Reflection Question** ❓ After solving Puzzle 2.1 and 2.2, Joe has the following claim:

"We know that the undergraduate courses are coded from 1 to 499, where a larger number (in the hundreds place) usually implies more advanced materials. Based on Puzzle 2.2, the results in Puzzle 2.1 clearly do not represent the hardest courses at Illinois."

Comment on Joe's claim below.

**Double-click and delete this line to type your comment.**

❓ **Group Discussion Question** ❓ Discuss Joe's claim with your group and whether or not you agree. How can we improve the analysis in Puzzle 2.1 or 2.2, e.g., using other descriptive statistics? 

## Puzzle 2.4: Easiest Courses in 2019

Based on the `"Average GPA"` column, find the **30 easiest courses in 2019** at Illinois. Store the results in `df_easy_course_2019`.

Hint (double-click this cell to see it): <span style="color:#ffffff00">use the nlargest() function</span>

<!-- grade_id: cell-44d17a2ef57b1a19, points: 0, grade: false, solution: true -->
```python
df_easy_course_2019 = ...
df_easy_course_2019
```

<!-- grade_id: cell-af77ae707b9fa776, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2.4 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error or output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, RE-RUN your code, and then RE-RUN this cell.
assert( len(df_easy_course_2019) == 30 ), "You should only have the 30 easiest courses in 2019."
assert( abs(df_easy_course_2019["Average GPA"].mean() - 3.9675) < 0.0001 ), "Are you sure you have the 30 easiest courses in 2019?"
## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

# Part 3 | Finding Data by Subject

## Puzzle 3.1: Total Number of Grades by Subject

Find the total number of each grade (eg: "A+"s, "A"s, "A-"s, etc.) given by different subjects. Store the results in `df_grouped_subject`.

Hint (double-click this cell to see it): <span style="color:#ffffff00">use the groupby() and agg() function</span>

<!-- grade_id: cell-7d31cf501bfd8025, points: 0, grade: false, solution: true -->
```python
df_grouped_subject = ...
df_grouped_subject
```

<!-- grade_id: cell-e9ad64dffa5a1cf1, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3.1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, RE-RUN your code, and then RE-RUN this cell.
assert( len(df_grouped_subject) == 161 ), "Are you sure you grouped by subjects?"
assert( df_grouped_subject[ df_grouped_subject["Subject"] == "CS" ]["A"].values[0] == 38472 ), "Are you sure you counted the right number of each subject?"
## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

## Puzzle 3.2: Average GPA by Subject

In `df_grouped_subject`, add a new column called `"Overall Average GPA"` that computes the overall GPA of each subject at Illinois.

<!-- grade_id: cell-7d3923eba40e1f45, points: 0, grade: false, solution: true -->
```python
df_grouped_subject["Overall Average GPA"] = ...
df_grouped_subject
```

<!-- grade_id: cell-81b547cf0d42b67e, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error or output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, RE-RUN your code, and then RE-RUN this cell.
assert( abs(df_grouped_subject[ df_grouped_subject["Subject"] == "CS" ]["Overall Average GPA"].values[0] \
    - 3.2494) < 0.0001 ), "Your overall average GPA does not appear to be correct."
## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

## Puzzle 3.3: Visualizing the Grades of Subject Equal to `"STAT"`

Plot a density histogram of the grades (`"A+"` to `"F"`) for the subject `"STAT"`.

Hint (double-click this cell to see it): 
<span style="color:#ffffff00">
    First, extract the row where the subject is "STAT".
    Then, extract the columns "A+" to "F" and divide them by the total number of students.
    Finally, use plot.bar().
</span>

```python

```

## Puzzle 3.4: Density Histogram vs. Frequency Histogram

❓ **Individual Reflection Question** ❓ Why should we use density histogram in Puzzle 3.3? Is it meaningful to compare the frequency histogram between different subjects?

**Double-click and delete this line to type your comment.**

❓ **Group Discussion Question** ❓ Think about what types of visualizations would be best for the GPA dataset.  These can be ones we have covered in lecture or other visualizations you are familiar with that we haven't covered in lecture. Discuss your ideas with your groupmates!

## Puzzle 3.5: Find Your Major!

Finally, find your major/subject in the dataset and check out how hard it is! 

Feel free to look around and explore! Part of being a data scientist is being creative and using Python to answer questions that you have.  Write your code below and write comments to explain what your code does. For example, what are the hardest classes in your major that have had at least 50 students?

<!-- grade_id: cell-116a2bb23ee81616, points: 0, grade: false, solution: true -->
```python
# Feel free to explore. Credit will be given as long as you try some new code (i.e., not directly copied from any puzzles).
```

❓ **Individual Reflection Question** ❓ Explain what you did in Puzzle 3.5 here, as well as what you found!  This is your chance to interpret your results and teach us something that you found interesting :)

**Double-click and delete this line to type your comment.**

❓ **Group Discussion Question** ❓ Are there any specific classes here at UIUC that you have heard of that are "hard" or "easy"?  Discuss this with your groupmates and check out their average GPAs. Think about how they compare to the overall average!

# Finale!

You're almost done -- congratulations!

You need to do two more things:

1.  Save your work. To do this, go to File -> Save All

2.  After you have saved, exit this notebook and follow the webpage instructions to commit this lab to your Git repository!